---
title: Test
description: Test
---

## Normal

```yaml
type: NormalExercise
key: 1be7e21a39
xp: 100
```

ok is this being deparsed? good save!

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
import math from pi
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## MPC

```yaml
type: MultipleChoiceExercise
key: c0ebadd267
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Pure MPC

```yaml
type: PureMultipleChoiceExercise
key: 3f52ddee45
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Tab

```yaml
type: TabExercise
key: fb63943e58
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: NormalExercise
key: c6859f90b6
xp: 35
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 22f4db7a00
xp: 35
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: PureMultipleChoiceExercise
key: d5a942f4da
xp: 30
```

---

## Bullet

```yaml
type: BulletExercise
key: 0d1198e4bb
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 1788feda6c
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: b3352e4624
xp: 50
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

---

## nice

```yaml
type: VideoExercise
key: 345e6ba6fb
xp: 50
```

`@projector_key`
3d1eed868536617107ebd8064a26b0de

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 894faf2813
xp: 100
```

normal ex now?

`@instructions`
instructions

`@hint`
hint

`@pre_exercise_code`
```{python}
import math from pi
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print(pi)
```

`@sct`
```{python}

```

---

## last

```yaml
type: NormalExercise
key: dcd32dac67
xp: 100
```

changed?

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
import math from pi


```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
